---
title: 'Template Chapter 1'
description: 'This is a template chapter.'
---

## An exercise title written in sentence case

```yaml
type: NormalExercise 
lang: r
xp: 100 
skills: 1
key: 513987f9cd   
```


这是R语言数据分析的入门课程


`@instructions`
- Instruction 1
- Instruction 2
- Instruction 3

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the way to the correct answer.
- So don't provide the answer, but don't just reiterate the instructions.
- Typically one hint per instruction is a sensible amount.

`@pre_exercise_code`

```{r}
# Load datasets and packages here.
```


`@sample_code`

```{r}
#请计算2+5

```


`@solution`

```{r}
2+5
```


`@sct`

```{r}
# Update this to something more informative.
success_msg("恭喜你成功书写你的第一条R语言代码")
```


